<a href="https://colab.research.google.com/github/Placide250/NLP/blob/main/n_placide_eagles_oc_edu.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

OPTIONAL: To test your solution against a subset of our test cases, you can [install the `otter` library](https://otter-grader.readthedocs.io/en/latest/index.html#installation) in your local python environment with your choice of the commands below and run the next cell:
```
pipx install otter-grader
# or
pip install otter-grader
```

```python
install.packages("ottr")
```

Output:
```
NameError                                 Traceback (most recent call last)
/tmp/ipython-input-3783777775.py in <cell line: 0>()
----> 1 install.packages("ottr")

NameError: name 'install' is not defined
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("shot_finder.ipynb")
# NOTE: If you change the name of this file, you will also need to update the line above to use Otter to test your solution.
```

**Question:** Given a list of shot dictionaries, implement `find_faulty_shots` to return a list of unique shotIDs for shots labeled incorrectly, in order of most recent to oldest.

Each shot dictionary contains the following fields:
- shotID: str
- period: int
- shotClockSeconds: float
- gameClockSeconds: float
- fg2Attempted: int
- fg2Made: int
- fg3Attempted: int
- fg3Made: int
- ftAttempted: int
- ftMade: int
- points: int

All field values are guaranteed to be correct except for shotClockSeconds, gameClockSeconds, and points.

A shot is labeled incorrectly if any of the following are true:
- It has an invalid shot clock value.
- It has an invalid game clock value. Note that gameClockSeconds represents the number of seconds remaining in the current period. A period value of 1-4 represents regulation and a period value of 5 or more represents overtime.
- It has an invalid points value, relative to the points expected based on shot type/zone, whether the shot was made, and the inclusion of free throws.

Each shotID is structured as MMDDYYYY_GG_FFFFFFFF, where:
- MMDDYYYY = date of the game
- GG = order in which the game was played relative to other games that day (e.g. 01, 02). You can assume that 02 starts after 01 finishes.
- FFFFFFFF = frame number of the shot within that game, left-padded to 8 digits (e.g. 00000001 for the first shot).

You may assume standard NBA rules for the shot clock, game clock, and point values.

```python
from datetime import datetime
# DO NOT IMPORT ANY ADDITIONAL LIBRARIES

def find_faulty_shots(shots_array: list[dict]) -> list[str]:
	# Replace the line below with your code
    ...
```

The cell below will run a subset of our tests using the `grader` generated in the cells above your implementation, but it will only work if you've installed the `otter` library.

```python
grader.check("q1")
```

# Task
Identify and list the unique `shotID`s of faulty shots from the provided list of shots, sorted in descending order by date. A shot is considered faulty if it violates standard NBA rules regarding shot clock, game clock, or point values.

## Understand shot rules

### Subtask:
Review the provided information about valid shot clock, game clock, and point values based on standard NBA rules.


## Define validation functions

### Subtask:
Create helper functions to validate each aspect of a shot: shot clock, game clock, and points.


**Reasoning**:
Define the helper functions to validate shot clock, game clock, and points based on standard NBA rules.



```python
def is_valid_shot_clock(shotClockSeconds: float) -> bool:
    """Validates the shot clock value."""
    return 0.0 <= shotClockSeconds <= 24.0

def is_valid_game_clock(period: int, gameClockSeconds: float) -> bool:
    """Validates the game clock value based on period type."""
    if period >= 1 and period <= 4:  # Regulation periods (12 minutes)
        return 0.0 <= gameClockSeconds <= 12 * 60
    elif period >= 5:  # Overtime periods (5 minutes)
        return 0.0 <= gameClockSeconds <= 5 * 60
    return False # Invalid period

def is_valid_points(fg2Attempted: int, fg2Made: int, fg3Attempted: int, fg3Made: int, ftAttempted: int, ftMade: int, points: int) -> bool:
    """Validates the points value based on shot and free throw attempts and makes."""
    expected_points = (fg2Made * 2) + (fg3Made * 3) + (ftMade * 1)
    return points == expected_points
```

## Iterate through shots

### Subtask:
Loop through the input list of shot dictionaries.


**Reasoning**:
Initialize an empty list and iterate through the input list of dictionaries as instructed.



```python
faulty_shot_ids = []
for shot in shots_array:
    # The logic to check for faulty shots will be added in the next step
    pass
```

Output:
```
NameError                                 Traceback (most recent call last)
/tmp/ipython-input-2106246954.py in <cell line: 0>()
      1 faulty_shot_ids = []
----> 2 for shot in shots_array:
      3     # The logic to check for faulty shots will be added in the next step
      4     pass

NameError: name 'shots_array' is not defined
```

**Reasoning**:
The error indicates that `shots_array` is not defined. This is because the function `find_faulty_shots` takes `shots_array` as an argument, but the code is being executed outside the function scope. The next step should be to put the initialization and iteration inside the function definition.



```python
from datetime import datetime
# DO NOT IMPORT ANY ADDITIONAL LIBRARIES

def find_faulty_shots(shots_array: list[dict]) -> list[str]:
    """Finds unique shot IDs of incorrectly labeled shots, sorted by date."""
    faulty_shot_ids = []
    for shot in shots_array:
        # The logic to check for faulty shots will be added in the next step
        pass
    # The sorting logic will be added in a later step
    return faulty_shot_ids
```

## Identify faulty shots

### Subtask:
For each shot, use the validation functions to check if it's faulty. If it is, add its `shotID` to a list of faulty shot IDs.


**Reasoning**:
Implement the logic inside the loop to check for faulty shots using the validation functions and append the shotID if faulty.



```python
def find_faulty_shots(shots_array: list[dict]) -> list[str]:
    """Finds unique shot IDs of incorrectly labeled shots, sorted by date."""
    faulty_shot_ids = []
    for shot in shots_array:
        if not is_valid_shot_clock(shot['shotClockSeconds']) or \
           not is_valid_game_clock(shot['period'], shot['gameClockSeconds']) or \
           not is_valid_points(shot['fg2Attempted'], shot['fg2Made'], shot['fg3Attempted'], shot['fg3Made'], shot['ftAttempted'], shot['ftMade'], shot['points']):
            faulty_shot_ids.append(shot['shotID'])
    # The sorting logic will be added in a later step
    return faulty_shot_ids
```

## Extract date from shotid

### Subtask:
Create a helper function to extract the date from a `shotID` string.


**Reasoning**:
Define the helper function `extract_date_from_shot_id` as instructed, which takes a shotID string, extracts the date part (first 8 characters), converts it to a datetime object, and returns it.



```python
from datetime import datetime

def extract_date_from_shot_id(shot_id: str) -> datetime:
    """Extracts the date from a shotID string and returns a datetime object."""
    date_str = shot_id[:8]  # Extract the first 8 characters (MMDDYYYY)
    return datetime.strptime(date_str, '%m%d%Y')
```

## Sort faulty shots

### Subtask:
Sort the list of faulty shot IDs in descending order based on the date extracted from the `shotID`.


**Reasoning**:
Sort the list of faulty shot IDs in descending order by date and ensure uniqueness.



```python
from datetime import datetime

def find_faulty_shots(shots_array: list[dict]) -> list[str]:
    """Finds unique shot IDs of incorrectly labeled shots, sorted by date."""
    faulty_shot_ids = []
    for shot in shots_array:
        if not is_valid_shot_clock(shot['shotClockSeconds']) or \
           not is_valid_game_clock(shot['period'], shot['gameClockSeconds']) or \
           not is_valid_points(shot['fg2Attempted'], shot['fg2Made'], shot['fg3Attempted'], shot['fg3Made'], shot['ftAttempted'], shot['ftMade'], shot['points']):
            faulty_shot_ids.append(shot['shotID'])

    # Sort by date in descending order and ensure uniqueness
    unique_faulty_shot_ids = sorted(list(set(faulty_shot_ids)), key=extract_date_from_shot_id, reverse=True)

    return unique_faulty_shot_ids
```

## Return faulty shots

### Subtask:
Return the sorted list of unique faulty `shotID`s.


## Summary:

### Data Analysis Key Findings

*   Helper functions were created to validate shot clock values (between 0.0 and 24.0 seconds), game clock values based on period type (0.0 to 12 minutes for periods 1-4, 0.0 to 5 minutes for periods 5+), and points values (calculated from made 2-point, 3-point, and free throw attempts).
*   A function `find_faulty_shots` was developed to iterate through a list of shot dictionaries and identify shots that violate any of the defined validation rules.
*   A helper function `extract_date_from_shot_id` was successfully implemented to parse the date from the first 8 characters of the `shotID` string in 'MMDDYYYY' format.
*   The identified faulty shot IDs were sorted in descending order based on the date extracted from their `shotID`s, and duplicate IDs were removed to ensure a unique list.

### Insights or Next Steps

*   The current analysis identifies shots that are potentially faulty based on rule violations. Further investigation into the specific reasons for these violations could provide insights into data collection errors or unusual game events.


Tests for q1 (points: None)
```
>>> def test_case_1():
...     shot_data = [{'shotID': '10282024_05_00000120', 'period': 1, 'shotClockSeconds': 16.2, 'gameClockSeconds': 723, 'fg2Attempted': 1, 'fg2Made': 0, 'fg3Attempted': 0, 'fg3Made': 0, 'fgAttempted': 1, 'fgMade': 0, 'ftAttempted': 0, 'ftMade': 0, 'points': 0}, {'shotID': '10282024_04_00642880', 'period': 3, 'shotClockSeconds': 21, 'gameClockSeconds': 57, 'fg2Attempted': 0, 'fg2Made': 0, 'fg3Attempted': 1, 'fg3Made': 1, 'fgAttempted': 1, 'fgMade': 1, 'ftAttempted': 0, 'ftMade': 0, 'points': 3}]
...     faulty_shots = find_faulty_shots(shot_data)
...     assert faulty_shots == ['10282024_05_00000120']
>>> test_case_1()
```
```
>>> def test_case_2():
...     shot_data = [{'shotID': '11022024_01_00000088', 'period': 1, 'shotClockSeconds': 8, 'gameClockSeconds': -11, 'fg2Attempted': 0, 'fg2Made': 0, 'fg3Attempted': 1, 'fg3Made': 0, 'fgAttempted': 1, 'fgMade': 0, 'ftAttempted': 0, 'ftMade': 0, 'points': 0}, {'shotID': '11032024_01_00000168', 'period': 2, 'shotClockSeconds': 7, 'gameClockSeconds': 88, 'fg2Attempted': 0, 'fg2Made': 0, 'fg3Attempted': 0, 'fg3Made': 0, 'fgAttempted': 0, 'fgMade': 0, 'ftAttempted': 2, 'ftMade': 1, 'points': 1}]
...     faulty_shots = find_faulty_shots(shot_data)
...     assert faulty_shots == ['11022024_01_00000088']
>>> test_case_2()
```
```
>>> def test_case_3():
...     shot_data = [{'shotID': '11042024_01_00000328', 'period': 1, 'shotClockSeconds': 12, 'gameClockSeconds': 671, 'fg2Attempted': 0, 'fg2Made': 0, 'fg3Attempted': 1, 'fg3Made': 1, 'fgAttempted': 1, 'fgMade': 1, 'ftAttempted': 0, 'ftMade': 0, 'points': 3}, {'shotID': '11042024_01_00002416', 'period': 2, 'shotClockSeconds': 18, 'gameClockSeconds': 643, 'fg2Attempted': 0, 'fg2Made': 0, 'fg3Attempted': 1, 'fg3Made': 0, 'fgAttempted': 1, 'fgMade': 0, 'ftAttempted': 0, 'ftMade': 0, 'points': 3}, {'shotID': '11042024_01_00001640', 'period': 1, 'shotClockSeconds': 5, 'gameClockSeconds': 26, 'fg2Attempted': 1, 'fg2Made': 1, 'fg3Attempted': 0, 'fg3Made': 0, 'fgAttempted': 1, 'fgMade': 1, 'ftAttempted': 1, 'ftMade': 1, 'points': 2}]
...     faulty_shots = find_faulty_shots(shot_data)
...     assert faulty_shots == ['11042024_01_00002416', '11042024_01_00001640']
>>> test_case_3()
```
```
>>> def test_case_4():
...     shot_data = [{'shotID': '11062024_02_00160336', 'period': 4, 'shotClockSeconds': 12, 'gameClockSeconds': 671, 'fg2Attempted': 0, 'fg2Made': 0, 'fg3Attempted': 1, 'fg3Made': 1, 'fgAttempted': 1, 'fgMade': 1, 'ftAttempted': 1, 'ftMade': 1, 'points': 4}, {'shotID': '11072024_01_00002096', 'period': 1, 'shotClockSeconds': 27, 'gameClockSeconds': 133, 'fg2Attempted': 1, 'fg2Made': 0, 'fg3Attempted': 0, 'fg3Made': 0, 'fgAttempted': 1, 'fgMade': 0, 'ftAttempted': 0, 'ftMade': 0, 'points': 0}]
...     faulty_shots = find_faulty_shots(shot_data)
...     assert faulty_shots == ['11072024_01_00002096']
>>> test_case_4()
```
